---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: dabffa6f2c
```














---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise

xp: 

key: 332d73c55f
```

This is a assignment.

`@instructions`
These are instructions

`@hint`
This is a hint

`@pre_exercise_code`
```{}
const foo = 'bar';
```


`@sct`
```{}
const foo = 'bar';
```




